<a href="https://colab.research.google.com/github/WHU-Peter/COMP6248-Deep-Learning/blob/master/2_1_ForwardAD.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Part 1: Forward Mode Automatic Differentiation

Forward mode AD can simply be implemented by defining a class to represent [dual numbers](https://en.wikipedia.org/wiki/Dual_number) which hold the value and its derivative. The following skeleton defines a dual number and implements multiplication. 

__Tasks:__

- Addition (`__add__`) is incomplete - can you finish it? 
- Can you also implement division (`__truediv__`), subtraction (`__sub__`) and power (`__pow__`)?

<!-- grade_id: cell-a645b47036f56c2e, points: 0, grade: false, solution: true -->
```python
import math

class DualNumber:
    def __init__(self, value, dvalue):
        self.value = value
        self.dvalue = dvalue

    def __str__(self):
        return str(self.value) + " + " + str(self.dvalue) + "ε"

    def __mul__(self, other):
        return DualNumber(self.value * other.value,
            self.dvalue * other.value + other.dvalue * self.value)
    
    def __add__(self, other):
        return DualNumber(self.value + other.value, self.dvalue + other.dvalue)

    def __truediv__(self, other):
        return DualNumber(self.value/other.value, (self.dvalue*other.value - self.value*other.dvalue)/(other.value*other.value))

    def __sub__(self, other):
        return DualNumber(self.value - other.value, self.dvalue - other.dvalue)
    
    def __pow__(self, other):
        result = DualNumber(1, 0)
        for i in range(other.value):
          result = result.__mul__(self)
        return result;
```

<!-- grade_id: cell-7923d5c39d574cd7, points: 3, grade: true, solution: false -->
```python
# Tests
print(DualNumber(1,0) + DualNumber(1,0) / DualNumber(1,0) - DualNumber(1,0)**DualNumber(1,0))
```

Output:
```
1.0 + 0.0ε
```

## Implementing math functions

We also need to implement some core math functions. Here's the sine function for a dual number:

```python
def sin(x):
    return DualNumber(math.sin(x.value), math.cos(x.value)*x.dvalue)
```

__Task:__ can you implement the _cosine_ (`cos`), _tangent_ (`tan`), and _exponential_ (`exp`) functions in the code block below?

<!-- grade_id: cell-61b9462cb781a8e1, points: 0, grade: false, solution: true -->
```python
# TODO: implement additional math functions on dual numbers

def cos(x):
    return DualNumber(math.cos(x.value), -math.sin(x.value)*x.dvalue)

def tan(x):
    return DualNumber(math.tan(x.value), 1/math.cos(2 * x.value)*x.dvalue)

def exp(x):
    return DualNumber(math.exp(x.value), math.exp(x.value)*x.dvalue)
```

<!-- grade_id: cell-9055ec23e6e9ffb8, points: 3, grade: true, solution: false -->
```python
# Tests
assert cos(DualNumber(0,0)).value == 1
assert tan(DualNumber(0,0)).value == 0
assert exp(DualNumber(0,0)).value == 1

```

## Time to try it out

We're now in a position to try our implementation.

__Task:__ 

- Try running the following code to compute the value of the function $z=x\cdot y+sin(x)$ given $x=0.5$ and $y=4.2$, together with the derivative $\partial z/\partial x$ at that point. 

<!-- grade_id: cell-12f7c8fa42784946, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x = DualNumber(0.5, 1)
y = DualNumber(4.2, 0)
z = x*y + sin(x)
print(z.value, z.dvalue)
```

Output:
```
2.579425538604203 5.077582561890373
```

__Task__: Differentiate the above function with respect to $x$ and write the symbolic derivatives in the following box. Verify the result computed above is correct by plugging-in the values into your symbolic gradient expression.

<!-- grade_id: cell-1ae4d0a0801a142a, points: 2, grade: true, solution: true -->
> dz/dx = y + cos(x)

__Task:__ Now use the code block below to compute the derivative $\partial z/\partial y$ of the above expression (at the same point $x=0.5, y=4.2$ as above) and store the derivative in the variable `dzdy` (just the derivative, not the Dual Number). Verify by hand that the result is correct.

<!-- grade_id: cell-eab763bf1b218dc8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x = DualNumber(0.5, 0)
y = DualNumber(4.2, 1)
z = x*y + sin(x)
print(z.value, z.dvalue)
dzdy = z.dvalue
```

Output:
```
2.579425538604203 0.5
```

<!-- grade_id: cell-a2f9f0ca4b1ed763, points: 2, grade: true, solution: false -->
```python
#Tests
assert dzdy
assert type(dzdy) == float
```

__Task:__ Finally, use the code block below to experiment and test the other math functions and methods you created.

<!-- grade_id: cell-b71d4bb64928ee38, points: 0, grade: true, solution: true -->
```python
import math

x = DualNumber(math.pi, 1)
z = sin(x)
print(z.value, z.dvalue)

x = DualNumber(math.pi, 1)
z = cos(x)
print(z.value, z.dvalue)

x = DualNumber(math.pi, 1)
z = tan(x)
print(z.value, z.dvalue)

x = DualNumber(1, 1)
z = exp(x)
print(z.value, z.dvalue)
```

Output:
```
1.2246467991473532e-16 -1.0
-1.0 -1.2246467991473532e-16
-1.2246467991473532e-16 1.0
2.718281828459045 2.718281828459045
```